## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc, to develop a new version of the function with an improved interface (the interface is the arguments that are passed to the function) with fewer arguments. Add appropriate checks inside your function on the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
import math
def area(a, b, c):
    if not (len(a)==2 and len(b)==2 and len(c)==2):
        print("Error: Input format incorrect")
        return 0
    if a==b or b==c or c==a:
        print("Error: Same points detected")
        return 0
    area = abs(a[0]*(b[1]-c[1])+b[0]*(c[1]-a[1])+c[0]*(a[1]-b[1]))/2
    if area == 0:
        print("Error: Area is zero")
    return area

## tests ##
a = [0.0, 0.0]
b = [0.0, 2.0]
c = [3.0, 0.0]
A = area(a, b, c)
assert math.isclose(A, 3.0)
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that, given only the vertex coordinates, computes the area of a simple polygon. Test your function for some simple shapes. Pay careful attention to the range of any loops. **Test for polygons without a vertex at the origin.**

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
def area_of_polygon(*args):
    x = []
    y = []
    area = 0
    for i in range(len(args)):
        x.append(args[i][0])
        y.append(args[i][1])
    x.append(x[0]) # as the "n+1 th" vertex
    y.append(y[0]) # as the "n+1 th" vertex
    #print(x, y)
    for i in range(len(args)):
        area+=(x[i] * y[i+1] - x[i+1] * y[i])
    return abs(area)/2

# test

A = [1, 2]
B = [2, 3]
C = [0, 1]
assert area_of_polygon(A, B, C) == 0

A = [0.0, 0.0]
B = [0.0, 2.0]
C = [3.0, 2.0]
D = [3.0, 0.0]
assert area_of_polygon(A, B, C, D) == 6.0

A = [1.0, 1.0]
B = [1.0, 3.0]
C = [4.0, 3.0]
D = [4.0, 1.0]
assert area_of_polygon(A, B, C, D) == 6.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match. Print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    if not len(x)==len(y):
        print("MISMATCH")
        return 0
    return [x[i]+y[i] for i in range(len(x))]

# tests ##
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
# tests ##
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Optional (advanced) 

Try writing a one-line version that performs the `sum_vector` operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
def sum_vector(x, y): return list(sum(i) for i in zip(x, y)) if len(x)==len(y) else "MISMATCH"

# test
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
sum_vector(a, b)
```

Output:
```
[-2, 11.3, -20, 8]
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A Python dictionary from college abbreviation to name; and
1. A Python list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
colleges = {
    "Trinity": "T",
    "Christ's": "CHR",
    "Churchill": "CHU",
    "Clare": "CL",
    "Clare Hall": "CLH",
    "Corpus Christi": "CC",
    "Darwin": "DAR",
    "Downing": "DOW",
    "Emmanuel": "EM",
    "Fitzwilliam": "F",
    "Girton": "G",
    "Gonville and Caius": "CAI",
    "Homerton": "HO",
    "Hughes Hall": "HH",
    "Jesus": "JE",
    "King's": "K",
    "Lucy Cavendish": "LC",
    "Magdalene": "M",
    "Murray Edwards": "MUR",
    "Newnham": "N",
    "Pembroke": "PEM",
    "Peterhouse": "PET",
    "Queens'": "Q",
    "Robinson": "R",
    "St Catharine's": "CTH",
    "St Edmund's": "ED",
    #"St John's": "JN",
    "Selwyn": "SE",
    "Sidney Sussex": "SID",
    "Trinity Hall": "TH",
    "Wolfson": "W"
}

colleges_reverse = {}
abbrevs = []

for name,abbrev in colleges.items():
    colleges_reverse[abbrev] = name
    abbrevs.append(abbrev)

print(colleges_reverse, sorted(abbrevs), sep="\n")
```

Output:
```
{'T': 'Trinity', 'CHR': "Christ's", 'CHU': 'Churchill', 'CL': 'Clare', 'CLH': 'Clare Hall', 'CC': 'Corpus Christi', 'DAR': 'Darwin', 'DOW': 'Downing', 'EM': 'Emmanuel', 'F': 'Fitzwilliam', 'G': 'Girton', 'CAI': 'Gonville and Caius', 'HO': 'Homerton', 'HH': 'Hughes Hall', 'JE': 'Jesus', 'K': "King's", 'LC': 'Lucy Cavendish', 'M': 'Magdalene', 'MUR': 'Murray Edwards', 'N': 'Newnham', 'PEM': 'Pembroke', 'PET': 'Peterhouse', 'Q': "Queens'", 'R': 'Robinson', 'CTH': "St Catharine's", 'ED': "St Edmund's", 'SE': 'Selwyn', 'SID': 'Sidney Sussex', 'TH': 'Trinity Hall', 'W': 'Wolfson'}
['CAI', 'CC', 'CHR', 'CHU', 'CL', 'CLH', 'CTH', 'DAR', 'DOW', 'ED', 'EM', 'F', 'G', 'HH', 'HO', 'JE', 'K', 'LC', 'M', 'MUR', 'N', 'PEM', 'PET', 'Q', 'R', 'SE', 'SID', 'T', 'TH', 'W']
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
...
```