Make NumPy available:

```python
import numpy as np
```

## Exercise 07.1 (indexing and timing)

Create two very long NumPy 1D-arrays (vectors) `x` and `y` and sum the arrays using:

1. The NumPy addition syntax, `z = x + y`; and
2. A `for` loop that computes the sum entry-by-entry

Compare the time required for the two approaches for arrays of different lengths (use a very long array for 
the timing). The values of the array entries are not important for this test. Use `%time` to report the time.

*Hint:* To loop over an array using indices, try a construction like:

```python
x = np.ones(10)
y = np.ones(len(x))
for i in range(len(x)):
    print(x[i]*y[i])
```

#### (1) Add two arrays using built-in addition operator:

<!-- grade_id: cell-2c856e54f7c3340e, points: 0, grade: false, solution: true -->
```python
import numpy as np
x = np.random.rand(1000000)
y = np.random.rand(1000000)
%time z = np.add(x, y)
```

Output:
```
CPU times: user 2.26 ms, sys: 46.4 ms, total: 48.6 ms
Wall time: 49.4 ms
```

#### (2) Add two arrays using own implementation:

<!-- grade_id: cell-f34614f9f0068cc4, points: 0, grade: false, solution: true -->
```python
import numpy as np
x = np.random.rand(1000000).tolist()
y = np.random.rand(1000000).tolist()
%time z = [x[i] + y[i] for i in range(1000000)]
```

Output:
```
CPU times: user 73.4 ms, sys: 19.6 ms, total: 93 ms
Wall time: 94.3 ms
```

### Optional extension: just-in-time (JIT) compilation

You will see a large difference in the time required between your NumPy and 'plain' Python implementations. This is due to Python being an *interpreted* language as opposed to a *compiled* language. A way to speed up plain Python implementions is to convert the interpreted Python code into compiled code. A tool for doing this is [Numba](https://numba.pydata.org/).

Below is an example using Numba and JIT to accelerate a computation:

```python
!pip -q install numba 
import numba
import math

def compute_sine_native(x):
    z = np.zeros(len(x))
    for i in range(len(z)):
        z[i] = math.sin(x[i])
    return z

@numba.jit
def compute_sine_jit(x):
    z = np.zeros(len(x))
    for i in range(len(z)):
        z[i] = math.sin(x[i])
    return z
    
x = np.ones(10000000)
%time z = compute_sine_native(x)
compute_sine_jit(x)
%time z = compute_sine_jit(x)
```

Output:
```
/tmp/ipykernel_68555/3117823105.py:11: NumbaDeprecationWarning: The 'nopython' keyword argument was not supplied to the 'numba.jit' decorator. The implicit default value for this argument is currently False, but it will be changed to True in Numba 0.59.0. See https://numba.readthedocs.io/en/stable/reference/deprecation.html#deprecation-of-object-mode-fall-back-behaviour-when-using-jit for details.
  @numba.jit
NameError                                 Traceback (most recent call last)
/home/bchen/Downloads/Forked/PartIA-Computing-Michaelmas/Exercises/07 Exercises.ipynb Cell 10 line 1
     <a href='vscode-notebook-cell:/home/bchen/Downloads/Forked/PartIA-Computing-Michaelmas/Exercises/07%20Exercises.ipynb#X12sZmlsZQ%3D%3D?line=14'>15</a>         z[i] = math.sin(x[i])
     <a href='vscode-notebook-cell:/home/bchen/Downloads/Forked/PartIA-Computing-Michaelmas/Exercises/07%20Exercises.ipynb#X12sZmlsZQ%3D%3D?line=15'>16</a>     return z
---> <a href='vscode-notebook-cell:/home/bchen/Downloa
```

**Task:** Test if Numba can be used to accelerate your implementation that uses indexing to sum two arrays, and by how much.

```python
...
```

## Exercise 07.2 (member functions and slicing)

Anonymised scores (out of 60) for an examination are stored in a NumPy array. Write:

1. A function that takes a NumPy array of the raw scores and returns the scores as percentages, sorted from 
   lowest to highest (try using `scores.sort()`, where `scores` is a NumPy array holding the scores).
1. A function that returns the maximum, minimum and mean of the raw scores as a dictionary with the 
   keys '`min`', '`max`' and '`mean`'. Use the NumPy array functions `min()`, `max()` and `mean()` to do the 
   computation, e.g. `max = scores.max()`.  
   
   Design your function for the min, max and mean to optionally exclude the highest and lowest scores from the 
   computation of the min, max and mean. 
   
   *Hint:* sort the array of scores and use array slicing to exclude
   the first and the last entries.

Use the scores 
```python
scores = np.array([58.0, 35.0, 24.0, 42, 7.8])
```
to test your functions.

<!-- grade_id: cell-169ebae60810c6be, points: 0, grade: false, solution: true -->
```python
import numpy as np

def to_percentage_and_sort(scores):
    scores.sort()
    return scores * 5 / 3

def statistics(scores, exclude=False):
    scores.sort()
    if exclude:
        scores = scores[1:len(scores)-1]
    return {'min': scores.min(), 'max': scores.max(), 'mean': scores.mean()}

## tests ##

scores = np.array([58.0, 35.0, 24.0, 42, 7.8])
assert np.isclose(to_percentage_and_sort(scores), [ 13.0, 40.0, 58.33333333,  70.0, 96.66666667]).all()

s0 = statistics(scores)
assert np.isclose(s0["min"], 7.8)
assert np.isclose(s0["mean"], 33.36)
assert np.isclose(s0["max"], 58.0)

s1 = statistics(scores, True)
assert np.isclose(s1["min"], 24.0)
assert np.isclose(s1["mean"], 33.666666666666666667)
assert np.isclose(s1["max"], 42.0)
```

<!-- grade_id: cell-af0b6fd8a3cadb1a, points: 0, grade: true, solution: false -->
```python
## tests ##

scores = np.array([58.0, 35.0, 24.0, 42, 7.8])
assert np.isclose(to_percentage_and_sort(scores), [ 13.0, 40.0, 58.33333333,  70.0, 96.66666667]).all()

s0 = statistics(scores)
assert np.isclose(s0["min"], 7.8)
assert np.isclose(s0["mean"], 33.36)
assert np.isclose(s0["max"], 58.0)

s1 = statistics(scores, True)
assert np.isclose(s1["min"], 24.0)
assert np.isclose(s1["mean"], 33.666666666666666667)
assert np.isclose(s1["max"], 42.0)
```

## Exercise 07.3 (slicing)

For the two-dimensional array

```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])
print(A)
```

use array slicing for the below operations, printing the results to the screen to check. Try to use array slicing such that your code would still work if the dimensions of `A` were enlarged.

**Check your results carefully against hand computations.**

#### 1. Extract the third column as a 1D array

<!-- grade_id: cell-dfee6b0ed9343682, points: 0, grade: false, solution: true -->
```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])
B = A[0:len(A), 2]
print(B)
```

Output:
```
[  -2.43   54.7     3.64 -113.2 ]
```

#### 2. Extract the first two rows as a 2D sub-array

<!-- grade_id: cell-7bf2f9a8c67029f8, points: 0, grade: true, solution: true -->
```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])
C = A[0:2, 0:A.shape[1]]
print(C)
```

Output:
```
[[ 4.    7.   -2.43 67.1 ]
 [-4.   64.   54.7  -3.33]]
```

#### 3.  Extract the bottom-right $2 \times 2$ block as a 2D sub-array

<!-- grade_id: cell-5206fea47d246222, points: 0, grade: false, solution: true -->
```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])
D = A[A.shape[0]-2:A.shape[0], A.shape[1]-2:A.shape[1]]
print(D)
```

Output:
```
[[  4.11   1.  ]
 [323.22   0.  ]]
```

#### 4. Sum the last column

<!-- grade_id: cell-34f74988e47b9f87, points: 0, grade: false, solution: true -->
```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])

number = np.sum(A[0:A.shape[0], A.shape[1]-1])
print(number)
```

Output:
```
391.1
```

#### Compute transpose

Compute the transpose of `A` (search online to find the function/syntax to do this).

<!-- grade_id: cell-dd2999d2da8070f6, points: 0, grade: false, solution: true -->
```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])
E = A.T
print(E)
```

Output:
```
[[   4.     -4.      2.43    1.2 ]
 [   7.     64.     23.2     2.5 ]
 [  -2.43   54.7     3.64 -113.2 ]
 [  67.1    -3.33    4.11  323.22]]
```

## Exercise 07.4 (optional extension)

In a previous exercise you implemented the bisection algorithm to find approximate roots of a mathematical function. Use the SciPy bisection function `optimize.bisect` (https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.bisect.html) to find roots of the mathematical function that was used in the previous exercise. Compare the results computed by SciPy and your program from the earlier exercise, and compare the computational time (using `%time`).

<!-- grade_id: cell-b6580accbcb3c2da, points: 0, grade: false, solution: true -->
```python
from scipy import optimize

...
```